```python
# Initialize Otter
import otter
grader = otter.Notebook("notebook06.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Notebook 06: Charts and Summary Statistics

<img src="./search-1975707_1280.jpg" width=500px/>

* Photo by damirbelavic from Pixabay
* Notebook created by Lily Lum with modifications by Shawn Wiggins 
* Last updated: September 2025  

## Learning Objectives

1. Create histograms and boxplots.
2. Computing summary statistics.

## 1. A couple notes before you start

<div style="background-color:#000; padding:10px;">
</div>

* This file is yours to edit and change.
* If you want to refer back to your copy in the future, you can find it in our JupyterHub.
* To run a cell, use `shift` + `enter`.  

## 2. Review on Loading Data

<div style="background-color:#000; padding:10px;">
</div>

The first step is to load in the data, which we will create a dataframe (i.e., spreadsheet) and name it `test_df`.

### Exercise 01 📍

1. Import the pandas library. Enter the code in line 2.  
2. On line 4, create a variable called `filepath`, and assign it to `'Testdata.csv'`, which is the name of the CSV file containing the data.
3. Check the header to make sure that everything worked properly. Enter the code in line 8. Recall that the code is `test_df.head()`.

Then run the cell, you will see a preview of the dataframe.    

```python
# Import pandas.

# Assign filepath to the CSV 'Testdata.csv'
...
# name the data frame "test_df"
test_df = pd.read_csv(filepath)
# Enter the code to read the header of the data set.

```

```python
grader.check("exercise_01")
```

## 3. Creating a Histogram for a Test Dataset

<div style="background-color:#000; padding:10px;">
</div>

Even though `matplotlib` is quite popular, the plots that it creates aren't particularly attractive without a lot of additional coding.

`seaborn` to the rescue! `seaborn` uses `matplotlib` functions behind the scenes; in other words, it does the hard work for you, allowing users to much more easily and simply create attractive plots. To learn more about the types of plots that `seaborn` can create, we encourage checking out the [tutorial](http://seaborn.pydata.org/tutorial.html) and [example gallery](http://seaborn.pydata.org/examples/index.html).  

To use functions within the `seaborn` library, we need to import both `matplotlib` and `seaborn`.

Finally, two notes about importing these libraries:
* When importing `matplotlib`, you need to include an additional line of code that allows plots to be displayed directly inside your notebook. See sample code below.
* When importing `seaborn`, we'll use the abbreviation `sns`. See sample code below.  


### Exercise 02 📍

The code below creates a histogram using the CSV file `Testdata.csv` for the column named `score`.

Run the cell to see the histogram for the column named `score`.

```python
import matplotlib
%matplotlib inline
import seaborn as sns
score_hg = sns.histplot(x = 'score', data = test_df)
```

```python
grader.check("exercise_02")
```

## 4. Computing the Summary Statistics for a Test Dataset

<div style="background-color:#000; padding:10px;">
</div>

Python has several helpful functions for computing the summary statistics of a dataframe. The code below computes the mean and standard deviation of the dataframe `test_df` and the column `score`
* sample mean:  `test_df['score'].mean()`  
* sample SD:  `test_df['score'].std()`  

### Exercise 03 📍

1. On line 2, enter the code to compute the sample mean score and assign it to the variable `x_bar`.  
2. On line 6, enter the code to compute the sample SD score and assign it to the variable `sd`.  
3. On line 8, enter the code to print the sample SD computed on line 6. Use line 4 as a guide.

Then run the cell, you will see the sample mean and SD score for the Testdata.

```python
# enter the code to compute the sample mean score and assign it to the variable x_bar

# print the sample mean computed
print('mean score =', x_bar)
# enter the code to compute the sample SD score and assign it to the variable sd

# enter the code to print the sample SD score computed

```

```python
grader.check("exercise_03")
```

## 5. Creating a Histogram for a COVID Data Set  

<div style="background-color:#000; padding:10px;">
</div>

The first step is to load in the data, which we will create a dataframe and name it `covid_df`. Then, you will create a histogram for the column `death rate`.

### Exercise 04 📍

1. On line 6, assign the `filepath` name to the CSV file `'COVIDbycountry.csv'`. Replace `...` with `'COVIDbycountry.csv'`.
2. Name the dataframe as `covid_df`. Update line 8 (i.e., change `xxx_df` to `covid_df`).    
3. On line 10, update the column name from `???` to `death rate`, and the data frame name from `xxx_df` to `covid_df`.  Notice that when making a histogram, you'd need to specify the column name and data frame.

Then run the cell to see the histogram for `death rate`.

```python
import pandas as pd
import matplotlib
%matplotlib inline
import seaborn as sns
# Assign filepath to the CSV 'COVIDbycountry.csv'
filepath = ...
# name the dataframe as covid_df.
xxx_df = pd.read_csv(filepath)
# update line 10 to the column name 'death rate' and the dataframe name to 'covid_df'.
deathrate_hg = sns.histplot(x = '???', data = xxx_df)
```

```python
grader.check("exercise_04")
```

### Exercise 05 📍

Make a histogram with `binwidth` wider than the previous histogram.  You'll need to modify the code by adding `binwidth = 0.02` to line 13 on [Exercise 03](#Exercise-03). Putting it together you'll have the following code:  

```python
deathrate_hg = sns.histplot(x = 'death rate', data = covid_df, binwidth = 0.02)
```

On Line 2, modify the code to include `binwidth = 0.02`. Then run the cell to see the histogram for `death rate`. What do you notice?

```python
# insert bindwith = 0.02 to the line of code below, change 'score' to 'death rate and 'test_df' to 'covid_df'
deathrate_hg = sns.histplot(x = 'death rate', data = covid_df)
# the code below allows you to change the label for the x-axis
deathrate_hg.set(xlabel= "Death Rate by Country")
```

```python
grader.check("exercise_05")
```

## 6. Creating Boxplot for COVID Dataset

<div style="background-color:#000; padding:10px;">
</div>

To create boxplots in Pyhton, we'll need to import `matplotlib` and `seaborn`. The function to create a boxplot is similar to creating a histogram. It is:  

```python
 sns.boxplot(x = 'column name', data = dataframe)
```

 For our `COVID by country` dataset, it would be:

```python
 sns.boxplot(x = 'death rate', data = covid_df)
```

Note, we can also use Python to compute the following summary statistics for the boxplot. The code to find the median, Q1, and Q3 are:
* sample median:  `covid_df['death rate'].median()`
* Q1, first quartile: `covid_df['death rate'].quantile(0.25)`
* Q3, third quartile: `covid_df['death rate'].quantile(0.75)`

### Exercise 06 📍

<!-- BEGIN QUESTION -->

On Line 6, update the column name and data frame, and then run the cell below to see the boxplot for `death rate` in the `COVIDbycountry` file and its summary statistics.

```python
# import matplotlib and seaborn
import matplotlib
%matplotlib inline
import seaborn as sns
# creating a boxplot for the column name the data frame to make a boxplot for death rate for the COVIDbycountry dataset.
sns.boxplot(x = '???', data = xxx_df)

# compute the median, Q1, and Q3
covid_median = covid_df['death rate'].median()
covid_Q1 = covid_df['death rate'].quantile(0.25)
covid_Q3 = covid_df['death rate'].quantile(0.75)

# print the median, Q1, and Q3
print('covid median =', covid_median)
print('covid Q1 =', covid_Q1)
print('covid Q3 =', covid_Q3)
```

<!-- END QUESTION -->

### Exercise 07 📍

Now, compute the mean and SD of death rate for the `COVIDbycountry.csv` dataset. Use the code in Exercise 3 as a guide.

1. On line 2, enter the code to compute the sample mean death rate and assign it to the variable `covid_mean`.  
2. On line 6, enter the code to compute the sample SD and assign it to the variable `covid_sd`.  
3. On line 8, enter the code to print the sample SD computed on line 6.

Then run the cell, you will see the sample mean and SD death rate for the `COVIDbycountry.csv` dataset.

```python
# enter the code to compute the sample mean death rate and assign it to the variable covid_mean

# print the sample mean computed
print('covid mean', covid_mean)
# enter the code to compute the sample SD death rate and assign it to the variable covid_sd

# enter the code to print the sample SD computed

```

```python
grader.check("exercise_07")
```

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="#/">APEX STATS Data and Measurement</a> created by David Schuster.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for exercise_01 (points: None)
```
>>> isinstance(test_df, pd.DataFrame)
True
```
failure_message: "❌ test_df is not assigned to a DataFrame."
success_message: "✅ test_df is assigned to a DataFrame."



Tests for exercise_02 (points: None)
```
>>> 'matplotlib' in globals()
True
```
failure_message: "❌ You have not imported matplotlib."
success_message: "✅ You have imported matplotlib."

```
>>> 'sns' in globals()
True
```
failure_message: "❌ You have not imported seaborn as sns."
success_message: "✅ You have imported seaborn as sns."

```
>>> isinstance(score_hg, matplotlib.axes._axes.Axes)
True
```
failure_message: "❌ score_hg is not assigned to a Matplotlib graphic."
success_message: "✅ score_hg is assigned to a Matplotlib graphic."



Tests for exercise_03 (points: None)
```
>>> import numpy as np
>>> bool(np.isclose(x_bar, 3.5555555555555554))
True
```
failure_message: "❌ The value you assigned to x_bar does not seem correct."
success_message: "✅ The value you assigned to x_bar seems correct."

```
>>> import numpy as np
>>> bool(np.isclose(sd, 2.6837686895641673))
True
```
failure_message: "❌ The value you assigned to sd does not seem correct."
success_message: "✅ The value you assigned to sd seems correct."



Tests for exercise_04 (points: None)
```
>>> isinstance(covid_df, pd.DataFrame)
True
```
failure_message: "❌ covid_df is not assigned to a DataFrame."
success_message: "✅ covid_df is assigned to a DataFrame."



Tests for exercise_05 (points: None)
```
>>> bin_widths = [p.get_width() for p in deathrate_hg.patches]
>>> all((abs(w - 0.02) < 0.001 for w in bin_widths))
True
```
failure_message: "❌ You ddi not set the binwidth to 0.02."
success_message: "✅ You set the binwidth to 0.02."



Tests for exercise_07 (points: None)
```
>>> import numpy as np
>>> bool(np.isclose(covid_mean, 0.018474418604651165))
True
```
failure_message: "❌ The value you assigned to covid_mean does not seem correct."
success_message: "✅ The value you assigned to covid_mean seems correct."

```
>>> import numpy as np
>>> bool(np.isclose(covid_sd, 0.02043242901904577))
True
```
failure_message: "❌ The value you assigned to covid_sd does not seem correct."
success_message: "✅ The value you assigned to covid_sd seems correct."

